<a href="https://colab.research.google.com/github/mahmud-nobe/CS110_pcw/blob/master/CS110%20Pre%20class%20work%203.1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Md Mahmudunnobe"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 3.1

## Question 1. 
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    # setting initial maximum summation as negative infinity
    max_sum = -float('inf')

    # checking every possible pairs and taking the sum of corresponding subarray
    # if the sum is larger than max_sum, then store it in max_sum and store the indices
    for i in range(len(A)):
      pair_sum = 0
      for j in range(i,len(A)):
        pair_sum += A[j]
        if (pair_sum > max_sum):
          max_sum = pair_sum
          max_i = i
          max_j = j
    return (max_i, max_j, max_sum)
    raise NotImplementedError()
```

```python
A = [-2, -3, 4, -1, -2, 1]
bruteforce_max_subarray(A)
```

Output:
```
(2, 2, 4)
```

```python

```

## Question 2. 
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
No, I don't need to run it again.

We already know the maximum subarray (indices and values) of the previous array A.

So, as we add only one new term, we can only check the subarrays, where the last index is the newly added one and compare them with the previous maximum subarray. 

## Question 3. 

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
**FALSE**

The only extra number that comes in array B is the last number, 8. So, if the maximum subarray of A and B are different, then it will be because of the last index and new maximum-subarray must include the last index. Otherwise their maximum subarry will be the same.

## Question 4.
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Inputs:
    - x: a NON-EMPTY list of numbers (e.g., the array B in the first two questions above). 
        * If x has 1 element: returns the value of the element regardless of the value of mx
    - mx: the maximum subarray of x excluding its last element (i.e., compute the 
    maximum subarray of the input array x considering only its first len(x) - 1 elements)
    
    Output: The maximum subarray of the array x.
    
    For example, using the array B in question 2, the result of incremental_max_subarray(B, 4) 
    is 10 (10 = 8 + 1 - 2 -1 + 4). 
    """
    max_sum = mx
    pair_sum = 0
    for i in reversed(range(len(x))):
      pair_sum += x[i]
      if (pair_sum > max_sum):
        max_sum = pair_sum
    
    return max_sum

    raise NotImplementedError()
```

```python
B = [-2, -3, 4, -1, -2, 1, 8]
incremental_max_subarray(B, 4)
```

Output:
```
10
```

```python
B[:1]
print(incremental_max_subarray(B[:1], 0))
print(B[0])
```

Output:
```
0
-2
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
# assert(incremental_max_subarray(B[:1], 0) == B[0])
# the last assert will be false as the B[0] = -2, but the mx = 0, which is larger than B[0],
# so the output will be 0, not B[0] (which equals -2)
```

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A to produce the value of the maximum
    subarray of A.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: flaot, the sum of the maximum subarray of A
    """

    if (len(A)==1):
      return A[0]
    else:
      return incremental_max_subarray(A, max_subarray(A[:-1]))

    raise NotImplementedError()

```

```python
B = [-2, -3, 4, -1, -2, 1, 8]
max_subarray(B)
```

Output:
```
10
```

```python
C = [1,2,3,4,5,6,7,8]
max_subarray(C)
```

Output:
```
36
```

```python
import random
D = random.sample(range(-50,50),10)
print(D)
max_subarray(D)
```

Output:
```
[-35, -19, 21, -15, 24, -38, -2, 42, 7, -50]
49
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2.
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
In the divide-conquer approach, our recurrance function is,

$T(n) = 2(T(n/2)) + \Theta(n)$

which by solving with tree method, results in the complexisty of

$\Theta(n) = n\,lg\,n$

Now let's analyze our approach:

*incremental_max_subarray* has only one for loop of length of x and inside of the loop, everything is linear.
so, it's complexity is $O(n)$

Then, in *max_subarray* we are dividing the problem into n subproblems with each of complexity of T(n/n) = T(1).

So, the recurrance is,

$T(n) = n(T(1)) + O(n)$

So the complexity is $O(n)$

So it is more effecient

```python

```