<a href="https://colab.research.google.com/github/mahmud-nobe/CS110_pcw/blob/master/CS110_Pre_Class_Work_5_1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Md Mahmudunnobe"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 5.1

## Question 1.

### Question 1a.

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
import random
import numpy as np
def worst_sort(A):
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Input:
    - A: list of numerical values
    
    Output:
    - A: Sorted list
    """
    
    # initially the value of sorted is True
    sorted = True
    
    # go through every element, if any element is bigger then its next element,
    # the array is not sorted, break the loop and make sorted False
    for i in range(len(A)-1): ### Complexity O(n) - best case
      if (A[i] > A[i+1]):
        sorted = False
        break
    
    # while the array is not sorted, suffle the array randomly
    while (not sorted):  ### O(E), where E = expected value of the number of suffle needed for sorting
      np.random.shuffle(A)

      # go through every element, if any element is bigger then its next element,
      # the array is not sorted, break the loop and make sorted False
      sorted = True
      for i in range(len(A)-1): ### O(n) ### Together: O(E)*O(n)
        if (A[i] > A[i+1]):
          sorted = False
          break
    
    return A
    
    raise NotImplementedError()
```

```python
A = [4,3,2,1,0,-1,-2,-3,-4]
worst_sort(A)
```

Output:
```
[-4, -3, -2, -1, 0, 1, 2, 3, 4]
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b.
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Question 1c.

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Question 1d.

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### [Optional] Question 1e.

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2.

Approximate median finder. Given a list and $δ$ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the $(50-\tilde{δ}/2)$% and $(50+\tilde{δ}/2)$% percentiles, where $\tilde{δ}=100\timesδ$. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a.
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is within (50-tilde{delta}/2)% and 
    (50+\tilde{delta}/2)% percentiles.
    X% percentile of a list is defined to be the smallest number in the list that is as large as at least
    X% numbers in the list. 
    
    Input:
    - A: list of numerical values
    - num: the number of interest
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - isApproxMed: a boolean value indicating whether num is within the bound
    """

    sortedA = worst_sort(A)

    p_value_high = 0.5+(delta/2)
    index_high = int(p_value_high * len(A)) 

    p_value_low = 0.5-(delta/2)
    index_low = int(p_value_low * len(A)) 

    print(p_value_low, p_value_high, index_low, index_high)
    return (num >= sortedA[index_low] and num <= sortedA[index_high])
    raise NotImplementedError()
```

```python
check_approx_med([0,1], 0, .25)
```

Output:
```
0.375 0.625 0 1
True
```

```python
check_approx_med([0], 0, .25)
```

Output:
```
0.375 0.625 0 0
True
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0,1], 0, .25) == True)
assert(check_approx_med([0], 0, .25) == False)
```

Output:
```
0.375 0.625 0 1
0.375 0.625 0 0
AssertionError                            Traceback (most recent call last)
<ipython-input-19-aa1f56e94d4b> in <module>()
      1 assert(check_approx_med([0,1], 0, .25) == True)
----> 2 assert(check_approx_med([0], 0, .25) == False)

AssertionError:
```

### Question 2b.
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 50-delta/2% and 50+delta/2% percentiles
    
    Input:
    - A: list of numerical values
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - num: the approximated median, if it is found within 100//delta trials (this is one possible 
    maximum number of retries. While you are encouraged to play around with another limits, the 
    submitted work must use this number.)
    - None: finding failed, if nothing found within 100//delta trials

    Note: the 100//delta is chosen here as it is the expected number of trials for the first successful finding to occur (using geometric distribution)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Question 2c.

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with δ or N (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Question 2d.
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE