<a href="https://colab.research.google.com/github/mahmud-nobe/CS110_pcw/blob/master/CS110_Pre_Class_Work_13_1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 13.1

## Question 1.

Complete the following functions, following the algorithms in Cormen et al.

<!-- grade_id: cell-542b2dce56293aef, points: 0, grade: false, solution: true -->
```python
def lcs_length(x, y):
    """
    Computes the length of an LCS of strings x and y.
    
    Inputs:
    - x, y: strings
    
    Outputs:
    - c: a list of lists of ints OR a numpy array. c[i,j] contains the 
    length of a LCS of x[:i] and y[:j]
    - b: a list of lists of strings OR a numpy array, containing the information
    used for LCS reconstruction (See Cormen et al.) Use "N" (North), "Y" 
    (North West), and "W" (West) that correspond to the directions of the arrows 
    used in Cormen et al.
    """
    
    # list for lengths of LCS and the information
    b = np.zeros((len(x), len(y)), dtype=str)
    c = np.zeros((len(x), len(y)))
    

    for i in range(len(x)):
        for j in range(len(y)):
            if x[i] == y[j]:
                if(i == 0  or j == 0):
                    c[i,j] = 1
                else:
                    c[i,j] = c[i-1, j-1] + 1
                b[i,j] = 'Y'
            elif (c[i-1, j] >= c[i,j-1]):
                c[i,j] = c[i-1, j]
                b[i,j] = 'N'
            else:
                c[i,j] = c[i, j-1]
                b[i,j] = 'W'
    
    return c, b

    # raise NotImplementedError()
```

<!-- grade_id: cell-6b825bf900c2cf0b, points: 0, grade: false, solution: true -->
```python
def print_lcs(b,x,i,j):
    """
    Finds a LCS.
    
    Inputs:
    - b: a list of lists of strings OR a numpy array, returned by lcs_length
    - x: string, an input to lcs_length
    - i, j: ints. print_lcs(b,x,i,j) returns a lcs of x[:i] and y[:j], where y
    is an input to lcs_length.
    
    Outputs:
    - lcs: list of strings, representing a LCS of x and y
    - length: int, the length of the LCS
    
    You can choose to actually PRINT OUT the LCS or not using the print function.
    
    """
    
    length = 0
    lcs = []

    while(i >= 0 and j >= 0):

        if(b[i,j] == 'Y'):
            #print(lcs)
            lcs += x[i]
            length += 1
            i,j = i-1, j-1
            
        elif(b[i,j] == 'N'):
            #print(b[i,j], i-1, j)
            i,j = i-1, j
        else:
            i,j = i, j-1
            #print_lcs(b,x,i,j-1)

    lcs.reverse()
    
    return (lcs, length)

    raise NotImplementedError()
```

```python

x, y = 'ambgdec', 'aubyci'
c, b = lcs_length(x, y)
print_lcs(b,x,len(x)-1,len(y)-1)[0]
```

Output:
```
['a', 'b', 'c']
```

<!-- grade_id: cell-e06e8f09bb8a5063, points: 1, grade: true, solution: false -->
```python
import numpy as np
x, y = 'ambgdec', 'aubyci'
c, b = lcs_length(x, y)
assert(print_lcs(b,x,len(x)-1,len(y)-1)[0] == ['a', 'b', 'c'])
assert(print_lcs(b,x,len(x)-1,len(y)-1)[1] == 3)

x, y = 'xyqwsssazdesaqqf', 'xoppoypllzookjdef'
c, b = lcs_length(x, y)
assert(print_lcs(b,x,len(x)-1,len(y)-1)[0]  == ['x', 'y', 'z', 'd', 'e', 'f'])
assert(print_lcs(b,x,len(x)-1,len(y)-1)[1]  == 6)

```

## Question 2. (Adapted from Exercise 15-4.1 Cormen et al.)
Use the functions built in Question 1 to find the LCS of `'10010101'` and `'010110110'`. You should store the list that represents the LCS you found in a variable named `lcs_q2`

<!-- grade_id: cell-5d742a82dde1229f, points: 0, grade: true, solution: true -->
```python
x, y = '10010101', '010110110'
c, b = lcs_length(x, y)
lcs_q2 = print_lcs(b,x,len(x)-1,len(y)-1)[0]
lcs_q2
# raise NotImplementedError()
```

Output:
```
['1', '0', '0', '1', '1', '0']
```

## Question 3. (Adapted from Exercise 15-4.5 Cormen et al.)
Complete the following function, making use of `lcs_length` and `print_lcs`. 



<!-- grade_id: cell-429c944f34a54c91, points: 0, grade: true, solution: true -->
```python
def lmis(lst):
    """
    Finds the Longest Monotonically Increasing Subsequence (LMIS) of a list 
    (lst) of n numbers in O(n^2) time. Note that a monotonically increasing 
    sequence is a sequence of numbers such that: a_1 <= a_2 <= ... <= a_n .
    
    Inputs:
    - lst: a list of ints
    
    Outputs:
    - out_lst: a list of ints, a longest monotonically increasing subsequence
    of lst
    """
    x = [str(i) for i in lst]
    y = [str(i) for i in sorted(lst)]

    c, b = lcs_length(x, y)
    lmis_lst = print_lcs(b,x,len(x)-1,len(y)-1)[0]


    lmis_int_lst = []

    for i in range(len(lmis_lst)):
        if(lmis_lst[i] == '-'):
            lmis_int_lst += [int(lmis_lst[i] + lmis_lst[i-1])]
            del lmis_int_lst[i-1]
        else:
            lmis_int_lst += [int(lmis_lst[i])]

    return lmis_int_lst

    #raise NotImplementedError()
```

```python
lmis([5, 1, 0, 4, 2, 6, 7, 9])
```

Output:
```
[1, 4, 6, 7, 9]
```

```python
lmis([6, -1,4,5,5,7,-6,2])
```

Output:
```
[-1, 4, 5, 5, 7]
```

<!-- grade_id: cell-f8b218898fb43d19, points: 0, grade: true, solution: false -->
```python
assert(lmis([5, 1, 0, 4, 2, 6, 7, 9]) == [1, 4, 6, 7, 9])
assert(lmis([6, -1,4,5,5,7,-6,2]) == [-1, 4, 5, 5, 7])
```

```python

```