<a href="https://colab.research.google.com/github/mahmud-nobe/CS110_pcw/blob/master/CS110_Pre_Class_Work_5_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 5.2

## Question 1.
Using Figure 7.1 in Cormen et al. as a model, perform manually the partition process on the following list: A = [1,5,6,2,3,8,9,4,7]. You just need to specify the followings:
1. The array after the process is done.
2. The value of $i$ after the process is done.

<!-- grade_id: cell-7aa520f8af13679b, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2.
Code up a Python implementation of `partition(A, p, r)`, closely follow the pseudo-code in Cormen et al., p.172. Your function should return the index of the pivot in the array.

<!-- grade_id: cell-1ceb2600756c60ff, points: 0, grade: false, solution: true -->
```python
def partition(A,p,r):
    """
    Assume r<len(A) and p>=0
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a57b60117a7b82fb, points: 1, grade: true, solution: false -->
```python
A = [1,5,6,2,3,8,9,4,7]
assert(partition(A, 0, len(A)-1)==6)
```

## Question 3.

Code up your own Python implementation of `quicksort(A, p, r)`, using `partition(A,p,r)`.

<!-- grade_id: cell-8c39ebb8cd1aa83a, points: 0, grade: false, solution: true -->
```python
def quick_sort(A,p,r):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4f822430efd456ee, points: 1, grade: true, solution: false -->
```python
A = [0]
assert(quick_sort(A, 0, 0) == [0])
A = [3,1,2]
assert(quick_sort(A, 0, 2) == [1,2,3])
```

## Question 4. 
Explain (using experimental plots) the running time of `quick_sort` when: 
1. all elements of array A have the same value (e.g., [1,1,1])?
2. array A contains distinct elements sorted in decreasing order (e.g., [5,4,2,1])?


<!-- grade_id: cell-b58035dd5fa02329, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```