<a href="https://colab.research.google.com/github/mahmud-nobe/CS110_pcw/blob/master/CS110_pre_class_work_3_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Md Mahmudunnobe"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 3.2

## Question 1.
Given the array `H=[39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]`, perform the following operations:
1. Draw the corresponding binary tree of H. Is the binary tree a valid max heap? Explain your answer.
2. Using as a model the drawing examples illustrated in Figure 6.2 of Cormen et al.,  draw a step-by-step transformation of the array above into a valid max heap. 
3. Now that you have obtained a valid max heap, write out the corresponding array that stores the valid max-heap.

Use as many cells as you wish for this question.

<!-- grade_id: cell-b1f821b034b44619, points: 0, grade: true, solution: true -->
1. Corresponding binary tree of H (note that array H is donated by A in the picture)
![heap1](https://github.com/mahmud-nobe/CS110_pcw/raw/master/Heap_1.jpg)






This binary tree is not a valid max-heap because every parent is not larger then both of its child, in this binary tree. For example, the first parent is 39, which is lower than both of its child (85 and 85) 

2. Step-by-step Transformation to a max-heap:

![heap1](https://github.com/mahmud-nobe/CS110_pcw/raw/master/Heap_1.jpg)

![heap2](https://github.com/mahmud-nobe/CS110_pcw/raw/master/Heap_2.jpg)

![heap3](https://github.com/mahmud-nobe/CS110_pcw/raw/master/Heap_3.jpg)

![heap4](https://github.com/mahmud-nobe/CS110_pcw/raw/master/Heap_4.jpg)


3. Corresponding array that stores the valid max-heap is already shown in the above picture:

- H_valid_max-heap = [92, 85, 85, 76, 49, 30, 49, 16, 39, 15, 21, 7, 29, 31, 28]

## Question 2. 
Consider the following questions on the $MAX-HEAPIFY$ operation.
### Question 2a.

In the pseudocode of $MAX-HEAPIFY$ (Cormen et al., p.154, or you can view it [here](https://drive.google.com/open?id=1e_3jsX4-qQCfZXKMok_T6LPFh9FwtmT5)), what does A.heap-size mean and what is the idea behind the local variable largest? 


<!-- grade_id: cell-06106b81a909b9e6, points: 0, grade: true, solution: true -->
**A.heap-size:**

A.heap-size is the number of element of which we want to build the binary tree. Usually the heap-size could be equal to the array size and we represent the whole array as a binary tree. But sometimes we could want the heap size to be smaller than the array size, in which case, we are removing the extra elements of the original array from our binary tree. This strategy can he useful in many cases including sorting an array using the heap methods.

**largest:**

The local variable *largest* stores the index number of the largest value of the subtree. In *max-heapify* function, we are comparing the parent with its left and right child and in *largest* variable we are storing the index number of the largest value among these three.

### Question 2b.
The functions $LEFT(i)$ and $RIGHT(i)$, lines 1 and 2 in the $MAX-HEAPIFY$ pseudocode, return the array index of the left and right child, respectively, of a node in a binary tree. From reading Section 6.1, you know that the input to both functions is an integer number, $i$, which corresponds to the array index of the parent node in the array. Review Section 6.1 for more information. Write a Python implementation of the functions $LEFT(i)$ and $RIGHT(i)$ by filling in the cells below.

<!-- grade_id: cell-2efd4017321c8d48, points: 0, grade: false, solution: true -->
```python
def left(i):
    # As python starts counting from 0, the left child's index no will be an odd number
    # and it will be 1 larger than the twice of the parent's index
    return 2*i + 1 
    raise NotImplementedError()
```

<!-- grade_id: cell-507ddbab62e6a32c, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

<!-- grade_id: cell-b7e5ceeacaeaf1b8, points: 0, grade: false, solution: true -->
```python
def right(i):
    # As python starts counting from 0, the right child's index no will be an even number
    # and it will be 2 larger than the twice of the parent's index
    return 2*i + 2
    raise NotImplementedError()
```

<!-- grade_id: cell-c24867515653788a, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 2c.
Write a Python implementation of the MAX-HEAPIFY operation using the pseudocode above, and your newly written functions, `left` and `right`.

<!-- grade_id: cell-52289253a243341b, points: 0, grade: false, solution: true -->
```python
def heapify(heap, i):
    """
    Inputs:
    - heap: a list of floats. Assume that the heap size is the length of the heap.
    
    No output is needed. This function should modify (if necessary) heap in-place.
    """
    # determine the index of the both child, where parent's index is i
    l = left(i)
    r = right(i)

    # check if there is any left child exist (is it less than the heap length) and
    # compare if the value of left child is larger than the value of parent
    # if yes, store the left child index in a variable 'largest', otherwise store the parent's index
    if (l < len(heap) and heap[l] > heap[i] ):
      largest = l
    else:
      largest = i

    # check if there is any right child exist (is it less than the heap length) and
    # compare the right child value with the maximum of the parent and left child
    # if it is larger, store its index at 'largest'
    if (r < len(heap) and heap[r] > heap[largest] ):
      largest = r

    # if the largest value is any of the child, then swap the child value with the parent and
    # call a recusrsive heapify function for the subtree rooted at the newly replaced child index
    if (largest != i):
      A[i], A[largest] = A[largest], A[i]
      heapify(heap, largest)
    #return heap
    #raise NotImplementedError()
```

<!-- grade_id: cell-7f63745ff34b881b, points: 1, grade: true, solution: false -->
```python
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
heapify(A,0)
assert(A == [85, 49, 85, 16, 39, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28])

```

## Question 3. 
Next, write a Python implementation of the BUILD_MAX_HEAP operation using the pseudocode provided in Section 6.3 of Cormen et. al. Test your Python implementation using the array in problem 1, and make sure your Python codes produce a valid max heap.

<!-- grade_id: cell-9237da6268a5e3eb, points: 0, grade: false, solution: true -->
```python
def build_max_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid max heap, in-place.
    """
    # every elements larger than len(A)/2 are the leaves of the heap tree
    # so heapify every subtree from the one rooted at len(A)/2 downto the subtree rooted at the 1st index (i.e. the main root of the tree)
    for i in reversed(range(int(len(A)/2))):
      heapify(A,i)
    #raise NotImplementedError()
```

<!-- grade_id: cell-5dc147b837da3f7e, points: 1, grade: true, solution: false -->
```python
A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
build_max_heap(A)
assert(A == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1])
```

## Question 4. 

Lastly, write Python implementations of the $MIN-HEAPIFY$ and $BUILD-MIN-HEAP$ operations for a min heap data structure. You can use your $MAX-HEAPIFY$ and $BUILD-MAX-HEAP$ Python function as models, just remember that the latter two functions support operations for the max heap data structure. Test your Python implementation of $BUILD-MIN-HEAP$ using the array in problem 1, and make sure your Python codes produce a valid min heap. 

<!-- grade_id: cell-d86943f127876fc8, points: 0, grade: false, solution: true -->
```python
def min_heapify(heap, i):
    """
    Inputs:
    - heap: a list of floats. Assume that the heap size is the length of the heap.
    
    No output is needed. This function should modify (if necessary) heap in-place.
    """
    # determine the index of the both child, where parent's index is i
    l = left(i)
    r = right(i)

    # check if there is any left child exist (is it less than the heap length) and
    # compare if the value of left child is smaller than the value of parent
    # if yes, store the left child index in a variable 'smallest', otherwise store the parent's index
    if (l < len(heap) and heap[l] < heap[i] ):
      smallest = l
    else:
      smallest = i

    # check if there is any right child exist (is it less than the heap length) and
    # compare the right child value with the minimum of the parent and left child
    # if it is smaller, store its index at 'smallest'
    if (r < len(heap) and heap[r] < heap[smallest] ):
      smallest = r

    # if the smallest value is any of the child, then swap the child value with the parent and
    # call a recusrsive min_heapify function for the subtree rooted at the newly replaced child index
    if (smallest != i):
      A[i], A[smallest] = A[smallest], A[i]
      min_heapify(heap, smallest)    
    
    # raise NotImplementedError()
```

```python

```

<!-- grade_id: cell-6efa45e3956ad0b8, points: 0, grade: false, solution: true -->
```python
def build_min_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid min heap, in-place.
    """
    # every elements larger than len(A)/2 are the leaves of the heap tree
    # so min_heapify every subtree from the one rooted at len(A)/2 downto the subtree rooted at the 1st index (i.e. the main root of the tree)
    for i in reversed(range(int(len(A)/2))):
      min_heapify(A,i)

    #raise NotImplementedError()
```

```python
A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
print(A)
build_min_heap(A)
print(A)
```

Output:
```
[4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
[1, 2, 3, 4, 7, 9, 10, 14, 8, 16]
```

```python
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
print(A)
build_min_heap(A)
print(A)
```

Output:
```
[39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
[7, 15, 28, 16, 21, 29, 31, 92, 76, 49, 85, 30, 85, 39, 49]
```

<!-- grade_id: cell-7e64941402f819c2, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```